Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Lyubomira Dimitrova"
COLLABORATORS = "Maryna Charniuk, Dung Nguyen"
```

---

# Aufgabe 1 - Treaps (40 Punkte)

Bei selbstbalancierenden Bäumen versucht man, die Höhe des Baumes zu minimieren. Dies ist aber gar nicht das eigentliche Ziel der Optimierung: In Wirklichkeit will man die Zugriffszeit auf die Elemente minimieren. Wenn die Schlüssel mit sehr unterschiedlicher Häufigkeit abgefragt werden, ist ein balancierter Baum dafür nicht die optimale Lösung. Stattdessen will man häufige Schlüssel nahe der Wurzel des Baumes abspeichern, auch wenn seltene Schlüssel dadurch in tieferen Ebenen landen als beim balancierten Baum. Dies wird durch den sogenannten Treap erreicht, der die Eigenschaften von Suchbäumen und Heaps verbindet. Neben dem Schlüssel hat hier jedes Element auch eine Priorität, und Elemente mit hoher Priorität liegen nahe der Wurzel des Baumes. Dies erreicht man, indem jeder Teilbaum zwei Eigenschaften erfüllt:  

>(1) Sortierung nach Schlüsseln: Alle Elemente im linken Teilbaum haben kleinere Schlüssel als die Wurzel des Teilbaums, alle Elemente im rechten Teilbaum größere (Suchbaumbedingung).  

>(2) Sortierung nach Prioritäten: Kein Element im linken oder rechten Teilbaum hat höhere Priorität als die Wurzel des Teilbaums (Heap-Bedingung).  

Die Erfinder der Treap-Datenstruktur haben gezeigt, dass beide Eigenschaften gleichzeitig erfüllbar sind. Der grundlegende Einfügealgorithmus besteht aus den folgenden Schritten:  

>1. Füge ein neues Element entsprechend seines Schlüssels ein wie in einen unbalancierten Baum (vergleiche **insert()** aus Übung 5). Damit ist Bedingung (1) erfüllt.  

>2. Wenn die Priorität des neuen Elements höher ist als die seines Vaterknotens: Führe eine Rota-tion aus, die das neue Element eine Ebene nach oben bewegt. Wenn das neue Element das linke Kind ist, muss eine Rechtsrotation ausgeführt werden und umgekehrt.  

>3. Wiederhole Schritt 2 auf der nächsthöheren Ebene bis entweder Bedingung (2) erfüllt oder das neue Element zur Wurzel des Treaps geworden ist. Da Bedingung (1) invariant gegenüber Rotationen ist, entsteht dadurch immer ein gültiger Treap.  

Für die Festlegung der Prioritäten gibt es mehrere Möglichkeiten:  

>(A)	Wenn man konstante Prioritäten verwendet, entsteht ein unbalancierter Baum.  

>(B)	Wenn man Zufallszahlen verwendet, entsteht ein näherungsweise balancierter Baum.  

>(C)	Wenn man die Wahrscheinlichkeiten verwendet, mit der auf die Schlüssel zugegriffen wird, entsteht ein näherungsweise zugriffsoptimaler Baum.  

>(D)	Wenn man bei jedem Zugriff auf ein Element dessen Priorität inkrementiert (und den Baum umstrukturiert, falls Bedingung (2) nicht mehr erfüllt ist), passt sich der Baum automatisch an variierende Zugriffsmuster an, indem häufig benutzte Schlüssel nach oben wandern. Beim ersten Einfügen wird die Priorität mit 1 initialisiert.  

Lösen Sie folgende Aufgaben:  

a) **(4 Points)**  
Fügen Sie die folgenden Zeichen mit ihren Prioritäten in einen leeren Treap ein und beantworton Sie die Fragen:

>(A, 17), (F, 22), (Q, 15), (K, 25), (R, 14), (B, 35), (X, 27), (M, 24), (L, 42), (T, 30)

1. Geben Sie die Länge des Treap an (0 heißt, dass sich nur ein Knoten im Treip befindet)
2. Welches Zeichen ist die Wurzel des Treap
3. Wie viele Blätter hat der Treap
4. Welches Zeichen (falls vorhanden) ist der rechte Kind-Knoten von M

<!-- grade_id: cell-2, points: 4, grade: true, solution: true -->

               (L, 42)
               /     \
           (B, 35)   (T, 30)
           /   \       /   \
       (A,17)(K,25) (M,24) (X, 27)
              /         \
           (F, 22)   (Q, 15)
                          \
                        (R, 14)
                        
1. Die Länge des Treaps ist 4.
2. Die Wurzel ist (L, 42).
3. Der Treap hat 4 Blätter.
4. (Q, 15) ist der rechte Kind-Knoten von M.

b) **(2 Points)**  
Löschen Sie (L, 42) aus dem Treap Und beantworten Sie dieselben Fragen.

<!-- grade_id: cell-4, points: 2, grade: true, solution: true -->
                           
                           (B, 35)
                           /     \
                       (A, 17)  (T, 30)
                                /     \
                            (K, 24)  (X, 27)
                            /     \
                        (F, 22)  (M, 24)
                                     \
                                    (Q, 15)
                                       \
                                      (R, 14)
                                      
                                      
1. Die Länge des Treaps ist 5.
2. Die Wurzel ist (B, 35).
3. Der Treap hat 4 Blätter.
4. (Q, 15) ist der rechte Kind-Knoten von M.

c)	**(2 Punkte)**   
Was ist der Unterschied zwischen den Eigenschaften des Binären-Suchbaums und der den Min-Heap? Kann die Min-Heap Eigenschaft ausgenutzt werden um die Schlüssel eines Baums mit $n$ Knoten in $O\left(n\right)$ auszugeben? Begründen Sie ihre Antwort.

<!-- grade_id: cell-6, points: 2, grade: true, solution: true -->
Suchbaumbedingung: leftChild < parent < rightChild <br>
MinHeap-Bedingung: child > parent

Nein, die Min-Heap Eigenschaft kann nicht dafür benutzt werden, alle n Schlüssel in sortierter Reihenfolge in $O(n)$ auszugeben, weil die Relation zwischen den rechten und linken Kindern nicht definiert ist. Im Gegensatz dazu lassen sich die Schlüssel eines Binärbaums in $O(n)$ sortiert ausgeben (z.B. in-order Traversierung).


d)	**(6 Punkte)**   
Implementieren Sie die aus der Vorlesung bekannten Operationen  
  >*newroot = rotateLeft(rootnode)*  
  >*newroot = rotateRight(rootnode)*
  
**Hinweis**   
Durch eine Rotation ändern wir die Struktur der Zeiger. Dies ist eine lokale Operation im Suchbaum durch die die Binäre-Suchbaum-Eigencshaft erhalten bleibt. Wenn wir beispielsweise eine Linksrotation an Knoten $x$ durchführen nehmen wir an, dass das rechte Kindelement von $x$, $y$ nicht leer ist; $x$ kann jeder Knoten im Baum sein, dessen rechtes Kind-Element nicht leer ist. Die Linksrotation dreht sich um die Verbindung von $x$ zu $y$. Dadurch wird $y$ die neue Wurzel des Teilbaums wobei $x$ das linke Kindelement von $y$ wird und das linke Kindelement von $y$ das rechte Kindelement von $x$ wird.

```python
class Node:
    def __init__(self, key, priority):
        self.key = key
        self.priority = priority
        self.left = self.right = None
    
    # for test
    def left_key(self):
        return 0 if self.left is None else self.left.key
    def right_key(self):
        return 0 if self.right is None else self.right.key
```

<!-- grade_id: cell-9, points: 0, grade: false, solution: true -->
```python
def rotateLeft(rootnode):
    right_child = rootnode.right
    rootnode.right = right_child.left
    right_child.left = rootnode
    return right_child    # the right child becomes the new root

def rotateRight(rootnode):
    left_child = rootnode.left
    rootnode.left = left_child.right
    left_child.right = rootnode
    return left_child     # the left child becomes the new root
```

<!-- grade_id: cell-10, points: 2, grade: true, solution: false -->
```python
# --- test left rotation ---
n1 = Node(1,1)
n2 = Node(2,2)
n1.right = n2
assert n1.left_key() == 0
assert n1.right_key() == 2
assert n2.left_key() == 0
assert n2.right_key() == 0
rotateLeft(n1)
assert n1.left_key() == 0
assert n1.right_key() == 0
assert n2.left_key() == 1
assert n2.right_key() == 0
# --- test right rotation ---
n1 = Node(1,1)
n2 = Node(2,2)
n2.left = n1
assert n1.left_key() == 0
assert n1.right_key() == 0
assert n2.left_key() == 1
assert n2.right_key() == 0
rotateRight(n2)
assert n1.left_key() == 0
assert n1.right_key() == 2
assert n2.left_key() == 0
assert n2.right_key() == 0
```

<!-- grade_id: cell-11, points: 4, grade: true, solution: false -->
```python
# --- test left rotation ---
n1 = Node(1,1)
n2 = Node(2,2)
n3 = Node(3,3)
n4 = Node(4,4)
n5 = Node(5,5)
n2.left = n1
n2.right = n4
n4.left = n3
n4.right = n5
assert n1.left_key() == 0
assert n1.right_key() == 0
assert n2.left_key() == 1
assert n2.right_key() == 4
assert n3.left_key() == 0
assert n3.right_key() == 0
assert n4.left_key() == 3
assert n4.right_key() == 5
assert n5.left_key() == 0
assert n5.right_key() == 0
rotateLeft(n2)
assert n1.left_key() == 0
assert n1.right_key() == 0
assert n2.left_key() == 1
assert n2.right_key() == 3
assert n3.left_key() == 0
assert n3.right_key() == 0
assert n4.left_key() == 2
assert n4.right_key() == 5
assert n5.left_key() == 0
assert n5.right_key() == 0
# --- test right rotation ---
n1 = Node(1,1)
n2 = Node(2,2)
n3 = Node(3,3)
n4 = Node(4,4)
n5 = Node(5,5)
n2.left = n1
n2.right = n3
n4.left = n2
n4.right = n5
assert n1.left_key() == 0
assert n1.right_key() == 0
assert n2.left_key() == 1
assert n2.right_key() == 3
assert n3.left_key() == 0
assert n3.right_key() == 0
assert n4.left_key() == 2
assert n4.right_key() == 5
assert n5.left_key() == 0
assert n5.right_key() == 0
rotateRight(n4)
assert n1.left_key() == 0
assert n1.right_key() == 0
assert n2.left_key() == 1
assert n2.right_key() == 4
assert n3.left_key() == 0
assert n3.right_key() == 0
assert n4.left_key() == 3
assert n4.right_key() == 5
assert n5.left_key() == 0
assert n5.right_key() == 0
```

e)	**(24 Punkte)**   
Implementieren Sie die Klassen **RandomTreap** und **DynamicTreap**, die die Prioritäten nach Variante (B) bzw. (D) festlegen. Gehen Sie dabei von der Klasse **SearchTree** aus Übung 5 aus. Die Node-Klasse muss um ein Attribut priority erweitert werden, das value-Attribut kann dafür wegfallen, weil es in dieser Übung nicht benötigt wird. Die Funktion **insert()** realisiert die Treap-Semantik: Ist der key noch nicht im Treap enthalten, wird er nach der Suchbaumregel eingefügt und der Baum eventuell entsprechend der Prioritäten umstrukturiert. Ist der key bereits vorhanden, passiert beim RandomTreap nichts, während beim DynamicTreap die Prio-rität um Eins erhöht und der Baum gegebenenfalls umstrukturiert wird.   
  
**Bemerkung**  
Die Funktion `insert()` von `RandomTreap`, hat ein optionales Argument `priority`. Wird dieses Argument weggelassen wird die Priorität zufällig durch Pythons `random.random()` Funktion festgelegt.

<!-- grade_id: cell-13, points: 0, grade: false, solution: true -->
```python
import random
### For test purpose, please use random.random() to set the priority of each node.

class RandomTreap:
    def __init__(self):
        self.root = None
        
    def insert(self, key, priority=None):
        self.root = self._insert_helper(self.root, key, priority)

    def _insert_helper(self, node, key, priority):
        if node is None:
            return Node(key, priority) if priority else Node(key, random.random())
        else:
            if key < node.key:
                node.left = self._insert_helper(node.left, key, priority)
                if node.left.priority > node.priority:
                    node = rotateRight(node)
            elif key > node.key:
                node.right = self._insert_helper(node.right, key, priority)
                if node.right.priority > node.priority:
                    node = rotateLeft(node)   
            
        return node
    

    def remove(self, key):
        self.root = self._remove_helper(self.root, key)
       
    def _remove_helper(self, node, key):
        if node is None:
            raise KeyError(key)

        if key == node.key:
            if node.left is None and node.right is None:     # case 1: node has no children
                return None
            
            if node.left is None:   # case 2.1: node only has right child
                return self._remove_helper(rotateLeft(node), key)
            if node.right is None:  # case 2.2: node only has left child
                return self._remove_helper(rotateRight(node), key)
            
            # case 3: node has two children
            node.priority = -10000
            if node.left.priority < node.right.priority:
                node = self._remove_helper(rotateLeft(node), key)
            else:
                node = self._remove_helper(rotateRight(node), key)

        elif key < node.key:
            node.left = self._remove_helper(node.left, key)
                
        elif key > node.key:
            node.right = self._remove_helper(node.right, key)

        return node
        
# Traverse
def preOrderTraverse(root):
    if root is not None:
        print ('key', root.key, 'priority', root.priority)
        preOrderTraverse(root.left)
        preOrderTraverse(root.right)
            
def inOrderTraverse(root):
    if root is not None:
        inOrderTraverse(root.left)
        print ('key', root.key, 'priority', root.priority)
        inOrderTraverse(root.right)

def postOrderTraverse(root):
    if root is not None:
        postOrderTraverse(root.left)
        postOrderTraverse(root.right)
        print ('key', root.key, 'priority', root.priority)

```

```python
# Check
def check(root):
    if root is not None:
        if root.left is not None:
            assert root.key >= root.left.key, "------------------\n!!! Error !!! Not a treap: the BST condition is violated!\n------------------"
            assert root.priority >= root.left.priority, "------------------\n!!! Error !!! Not a treap: the heap condition is violated!\n------------------"
        if root.right is not None:
            assert root.key <= root.right.key, "------------------\n!!! Error !!! Not a treap: the BST condition is violated!\n------------------"
            assert root.priority >= root.right.priority, "------------------\n!!! Error !!! Not a treap: the heap condition is violated!\n------------------"
        check(root.left)
        check(root.right)
```

<!-- grade_id: cell-14, points: 4, grade: true, solution: false -->
```python
random.seed(1)
a = RandomTreap()
a.insert(1,17) 
a.insert(6,22)
a.insert(17,15)
a.insert(11,25)
a.insert(18,14)
a.insert(2,35)
a.insert(24,27)
a.insert(13,24)
a.insert(12,42)
a.insert(20,30)
print('--- preOrderTraverse ---')
preOrderTraverse(a.root)
print('--- inOrderTraverse ---')
inOrderTraverse(a.root)
print('--- postOrderTraverse ---')
postOrderTraverse(a.root)
check(a.root)
a.remove(12)
print('--- preOrderTraverse ---')
preOrderTraverse(a.root)
print('--- inOrderTraverse ---')
inOrderTraverse(a.root)
print('--- postOrderTraverse ---')
postOrderTraverse(a.root)
check(a.root)
```

Output:
```
--- preOrderTraverse ---
key 12 priority 42
key 2 priority 35
key 1 priority 17
key 11 priority 25
key 6 priority 22
key 20 priority 30
key 13 priority 24
key 17 priority 15
key 18 priority 14
key 24 priority 27
--- inOrderTraverse ---
key 1 priority 17
key 2 priority 35
key 6 priority 22
key 11 priority 25
key 12 priority 42
key 13 priority 24
key 17 priority 15
key 18 priority 14
key 20 priority 30
key 24 priority 27
--- postOrderTraverse ---
key 1 priority 17
key 6 priority 22
key 11 priority 25
key 2 priority 35
key 18 priority 14
key 17 priority 15
key 13 priority 24
key 24 priority 27
key 20 priority 30
key 12 priority 42
--- preOrderTraverse ---
key 2 priority 35
key 1 priority 17
key 20 priority 30
key 11 priority 25
key 6 priority 22
key 13 priority 24
key 17 priority 15
key 18 priority 14
key 24 priority 27
--- inOrderTraverse ---
key 1 priority 17
key 2 priority 35
key 6 priority 22
key 11 priority 25
key 13 priority 24
key 17 priority 15
key 18 priority 14
key 20 priority 3
```

<!-- grade_id: cell-15, points: 4, grade: true, solution: false -->
```python
random.seed(1)
a = RandomTreap()
a.insert(6) 
a.insert(2)
a.insert(5)
a.insert(3)
a.insert(8)
a.insert(4)
a.insert(1)
a.insert(7)
print('--- preOrderTraverse ---')
preOrderTraverse(a.root)
print('--- inOrderTraverse ---')
inOrderTraverse(a.root)
print('--- postOrderTraverse ---')
postOrderTraverse(a.root)
check(a.root)
```

Output:
```
--- preOrderTraverse ---
key 2 priority 0.8474337369372327
key 1 priority 0.651592972722763
key 7 priority 0.7887233511355132
key 5 priority 0.763774618976614
key 4 priority 0.4494910647887381
key 3 priority 0.2550690257394217
key 6 priority 0.13436424411240122
key 8 priority 0.49543508709194095
--- inOrderTraverse ---
key 1 priority 0.651592972722763
key 2 priority 0.8474337369372327
key 3 priority 0.2550690257394217
key 4 priority 0.4494910647887381
key 5 priority 0.763774618976614
key 6 priority 0.13436424411240122
key 7 priority 0.7887233511355132
key 8 priority 0.49543508709194095
--- postOrderTraverse ---
key 1 priority 0.651592972722763
key 3 priority 0.2550690257394217
key 4 priority 0.4494910647887381
key 6 priority 0.13436424411240122
key 5 priority 0.763774618976614
key 8 priority 0.49543508709194095
key 7 priority 0.7887233511355132
key 2 priority 0.8474337369372327
```

<!-- grade_id: cell-16, points: 2, grade: true, solution: false -->
```python
a.insert(8)
a.insert(4)
a.insert(1)
a.insert(7)
assert a.root.key == 2
assert a.root.left.key == 1
assert a.root.right.key == 7
assert a.root.right.left.key == 5
assert a.root.right.left.left.key == 4
assert a.root.right.left.left.left.key == 3
assert a.root.right.left.right.key == 6
assert a.root.right.right.key == 8
```

<!-- grade_id: cell-17, points: 4, grade: true, solution: false -->
```python
a.remove(7)
check(a.root)
a.remove(1)
check(a.root)
a.insert(1)
check(a.root)
a.insert(7)
check(a.root)
print('--- preOrderTraverse ---')
preOrderTraverse(a.root)
print('--- inOrderTraverse ---')
inOrderTraverse(a.root)
print('--- postOrderTraverse ---')
postOrderTraverse(a.root)
```

Output:
```
--- preOrderTraverse ---
key 2 priority 0.8474337369372327
key 1 priority 0.0938595867742349
key 5 priority 0.763774618976614
key 4 priority 0.4494910647887381
key 3 priority 0.2550690257394217
key 8 priority 0.49543508709194095
key 6 priority 0.13436424411240122
key 7 priority 0.02834747652200631
--- inOrderTraverse ---
key 1 priority 0.0938595867742349
key 2 priority 0.8474337369372327
key 3 priority 0.2550690257394217
key 4 priority 0.4494910647887381
key 5 priority 0.763774618976614
key 6 priority 0.13436424411240122
key 7 priority 0.02834747652200631
key 8 priority 0.49543508709194095
--- postOrderTraverse ---
key 1 priority 0.0938595867742349
key 3 priority 0.2550690257394217
key 4 priority 0.4494910647887381
key 7 priority 0.02834747652200631
key 6 priority 0.13436424411240122
key 8 priority 0.49543508709194095
key 5 priority 0.763774618976614
key 2 priority 0.8474337369372327
```

Implementieren Sie  `DynamicTreap` in der folgenden Zelle.

<!-- grade_id: cell-19, points: 0, grade: false, solution: true -->
```python
class DynamicTreap:
    def __init__(self):
        self.root = None

    def insert(self, key):
        self.root = self._insert_helper(self.root, key)

    def _insert_helper(self, node, key):
        if node is None:
            return Node(key, 1) 
        else:
            if key < node.key:
                node.left = self._insert_helper(node.left, key)
                if node.left.priority > node.priority:
                    node = rotateRight(node)
            elif key > node.key:
                node.right = self._insert_helper(node.right, key)
                if node.right.priority > node.priority:
                    node = rotateLeft(node)
            else:
                node.priority += 1   
        return node
    
    def remove(self, key):
        self.root = self._remove_helper(self.root, key)
       
    def _remove_helper(self, node, key):
        if node is None:
            raise KeyError(key)

        if key == node.key:
            if node.left is None and node.right is None:     # case 1: node has no children
                return None
            
            if node.left is None:   # case 2.1: node only has right child
                return self._remove_helper(rotateLeft(node), key)
            if node.right is None:  # case 2.2: node only has left child
                return self._remove_helper(rotateRight(node), key)
            
            # case 3: node has two children
            node.priority = -1000         # negative priority, so the max heap condition will push the node to remove to a leaf
            if node.left.priority < node.right.priority:
                node = self._remove_helper(rotateLeft(node), key)
            else:
                node = self._remove_helper(rotateRight(node), key)

        elif key < node.key:
            node.left = self._remove_helper(node.left, key)
                
        elif key > node.key:
            node.right = self._remove_helper(node.right, key)

        return node

```

<!-- grade_id: cell-20, points: 6, grade: true, solution: false -->
```python
b = DynamicTreap()
b.insert(6) 
b.insert(2)
b.insert(5)
b.insert(3)
b.insert(8)
b.insert(4)
b.insert(1)
b.insert(7)
print('--- preOrderTraverse ---')
preOrderTraverse(b.root)
print('--- inOrderTraverse ---')
inOrderTraverse(b.root)
print('--- postOrderTraverse ---')
postOrderTraverse(b.root)
check(b.root)
b.insert(5)
b.insert(4)
b.insert(7)
b.insert(5)
b.insert(7)
b.insert(7)
print('--- preOrderTraverse ---')
preOrderTraverse(b.root)
print('--- inOrderTraverse ---')
inOrderTraverse(b.root)
print('--- postOrderTraverse ---')
postOrderTraverse(b.root)
check(b.root)
```

Output:
```
--- preOrderTraverse ---
key 6 priority 1
key 2 priority 1
key 1 priority 1
key 5 priority 1
key 3 priority 1
key 4 priority 1
key 8 priority 1
key 7 priority 1
--- inOrderTraverse ---
key 1 priority 1
key 2 priority 1
key 3 priority 1
key 4 priority 1
key 5 priority 1
key 6 priority 1
key 7 priority 1
key 8 priority 1
--- postOrderTraverse ---
key 1 priority 1
key 4 priority 1
key 3 priority 1
key 5 priority 1
key 2 priority 1
key 7 priority 1
key 8 priority 1
key 6 priority 1
--- preOrderTraverse ---
key 7 priority 4
key 5 priority 3
key 4 priority 2
key 2 priority 1
key 1 priority 1
key 3 priority 1
key 6 priority 1
key 8 priority 1
--- inOrderTraverse ---
key 1 priority 1
key 2 priority 1
key 3 priority 1
key 4 priority 2
key 5 priority 3
key 6 priority 1
key 7 priority 4
key 8 priority 1
--- postOrderTraverse ---
key 1 priority 1
key 3 priority 1
key 2 priority 1
key 4 priority 2
key 6 priority 1
key 5 priority 3
key 8 priority 1
key 7 priority 4
```

<!-- grade_id: cell-21, points: 6, grade: true, solution: false -->
```python
b.remove(5)
print('--- preOrderTraverse ---')
preOrderTraverse(b.root)
print('--- inOrderTraverse ---')
inOrderTraverse(b.root)
print('--- postOrderTraverse ---')
postOrderTraverse(b.root)
check(b.root)
b.insert(5)
print('--- preOrderTraverse ---')
preOrderTraverse(b.root)
print('--- inOrderTraverse ---')
inOrderTraverse(b.root)
print('--- postOrderTraverse ---')
postOrderTraverse(b.root)
check(b.root)
```

Output:
```
--- preOrderTraverse ---
key 7 priority 4
key 4 priority 2
key 2 priority 1
key 1 priority 1
key 3 priority 1
key 6 priority 1
key 8 priority 1
--- inOrderTraverse ---
key 1 priority 1
key 2 priority 1
key 3 priority 1
key 4 priority 2
key 6 priority 1
key 7 priority 4
key 8 priority 1
--- postOrderTraverse ---
key 1 priority 1
key 3 priority 1
key 2 priority 1
key 6 priority 1
key 4 priority 2
key 8 priority 1
key 7 priority 4
--- preOrderTraverse ---
key 7 priority 4
key 4 priority 2
key 2 priority 1
key 1 priority 1
key 3 priority 1
key 6 priority 1
key 5 priority 1
key 8 priority 1
--- inOrderTraverse ---
key 1 priority 1
key 2 priority 1
key 3 priority 1
key 4 priority 2
key 5 priority 1
key 6 priority 1
key 7 priority 4
key 8 priority 1
--- postOrderTraverse ---
key 1 priority 1
key 3 priority 1
key 2 priority 1
key 5 priority 1
key 6 priority 1
key 4 priority 2
key 8 priority 1
key 7 priority 4
```