Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Lyubomira Dimitrova"
COLLABORATORS = "Maryna Charniuk, Dung Nguyen"
```

---

# Aufgabe 1 - Binäre Suchbäume (16 Punkte)

a)	**(5 Punkte)**  
In der Vorlesung haben wir die Klasse SearchTree behandelt, die (key, value)-Paare in Node-Objekten mit folgender Definition ablegt:

```python
class Node:
    def __init__(self, key, value):
        self.key = key
        self.value = value
        self.left = self.right = None
```

Das Grundgerüst der Klasse SearchTree hat folgende Form:

<!-- grade_id: task1a-class2, points: 0, grade: false, solution: true -->
```python
class SearchTree:
    def __init__(self):
        self.root = None
        self.size = 0

    def __len__(self):
        return self.size
    
    def insert(self, key, value):
        self.root = self._insert_helper(self.root, key, value)

    def _insert_helper(self, node, key, value):
        if node is None:
            self.size += 1
            return Node(key, value)
        else:
            if key < node.key:
                node.left = self._insert_helper(node.left, key, value)
                return node
            elif key > node.key:
                node.right = self._insert_helper(node.right, key, value)
                return node
            else:
                node.key = key
                node.value = value
                return node
            
        
    def remove(self, key):
        self.root = self._remove_helper(self.root, key)
        self.size -= 1

    def _remove_helper(self, node, key):
        if node is None:
            raise KeyError(key)

        if key == node.key:
            if node.left is None and node.right is None:     # case 1: node has no children
                return None
            if node.left is None:   # case 2.1: node only has right child
                return node.right
            if node.right is None:  # case 2.2: node only has left child
                return node.left
            
            # case 3: node has two children
            replacement = node.left      
            while replacement.right is not None:    # find largest key in left sub-tree
                replacement = replacement.right

            node.key = replacement.key
            node.value = replacement.value
            node.left = self._remove_helper(node.left, replacement.key)   # delete replacement node from 
            return node                                                   # its spot in the left subtree

        elif key < node.key:
            node.left = self._remove_helper(node.left, key)
            return node
        elif key > node.key:
            node.right = self._remove_helper(node.right, key)
            return node
        
        
    def find(self, key):
        current = self.root

        while current is not None and current.key != key:
            if current.key > key:
                current = current.left
            else:
                current = current.right
        return current
    

        
# Traverse
def preOrderTraverse(root):
    if root is not None:
        print ('key', root.key, 'value', root.value)
        preOrderTraverse(root.left)
        preOrderTraverse(root.right)
            
def inOrderTraverse(root):
    if root is not None:
        inOrderTraverse(root.left)
        print ('key', root.key, 'value', root.value)
        inOrderTraverse(root.right)

def postOrderTraverse(root):
    if root is not None:
        postOrderTraverse(root.left)
        postOrderTraverse(root.right)
        print ('key', root.key, 'value', root.value)

```

**Hint:**  
Geben Sie einer vollständigen Implementation ab. Wenn tree ein Objekt vom Typ SearchTree ist, soll die Semantik der Aufrufe wie folgt realisiert werden:  

>*tree.insert(key, value):*   
Fügt den gegebenen Wert value unter dem Schlüssel key ein. Falls key schon vorhanden war, soll der zuvor gespeicherte Wert überschrieben werden. Andernfalls soll ein neuer Node so eingefügt werden, dass die Suchbaumbedingung erhalten bleibt.  

>*tree.remove(key):*   
Löscht den Node, der den gegebenen Schlüssel enthält, wobei die Such-baumbedingung erhalten bleibt. Falls der Schlüssel nicht vorhanden ist, soll eine KeyError-Excpetion ausgelöst werden, deren Fehlermeldung den Schlüssel angibt.  

>*found=tree.find(key):*   
Gibt den Node zurück, der den Schlüssel key enthält, oder None, wenn der Schlüssel nicht im Baum vorhanden ist.


<!-- grade_id: task1a-test, points: 5, grade: true, solution: false -->
```python
a = SearchTree();
assert a.__len__() == 0
a.insert(6,60) 
a.insert(2,20)
a.insert(5,50)
a.insert(3,30)
a.insert(8,80)
a.insert(4,40)
a.insert(1,10)
a.insert(7,70)
assert a.root.left.right.left.right.key == 4
assert a.root.left.right.left.right.value == 40
assert a.root.right.left.key == 7
assert a.root.right.left.value == 70
node1 = a.find(5)
assert node1.key == 5
assert node1.value == 50
a.insert(5,500)
node1 = a.find(5)
assert node1.key == 5
assert node1.value == 500
a.remove(5)
assert a.find(5) == None

print('preOrderTraverse:')
preOrderTraverse(a.root)
print('inOrderTraverse:')
inOrderTraverse(a.root)
print('postOrderTraverse:')
postOrderTraverse(a.root)
    
```

Output:
```
preOrderTraverse:
key 6 value 60
key 2 value 20
key 1 value 10
key 3 value 30
key 4 value 40
key 8 value 80
key 7 value 70
inOrderTraverse:
key 1 value 10
key 2 value 20
key 3 value 30
key 4 value 40
key 6 value 60
key 7 value 70
key 8 value 80
postOrderTraverse:
key 1 value 10
key 4 value 40
key 3 value 30
key 2 value 20
key 7 value 70
key 8 value 80
key 6 value 60
```

b)	**(4 Punkte)**  
Entwickeln Sie einen Algorithmus, der die Tiefe des Baumes (den maximalen Abstand von der Wurzel zu einem Blatt) bestimmt und implementieren Sie ihn als Methode, so dass depth=tree.depth() die Tiefe zurückgibt.   
**Hint:** If there is only one node, its depth is 0. 

<!-- grade_id: task1b-answer, points: 0, grade: false, solution: true -->
```python
def depth(tree):
    def max_depth(node):
        if node is None:
            return -1
        else:
            max_d_left = max_depth(node.left)
            max_d_right = max_depth(node.right)
            max_d = max(max_d_left, max_d_right) + 1
            return max_d
    
    return max_depth(tree.root)

SearchTree.depth = depth
```

<!-- grade_id: task1b-test, points: 4, grade: true, solution: false -->
```python
b = SearchTree();
assert b.depth() == -1
b.insert(6,60)
assert b.depth() == 0
b.insert(2,20)
assert b.depth() == 1
b.insert(5,50)
assert b.depth() == 2
b.insert(3,30)
assert b.depth() == 3
b.insert(8,80)
assert b.depth() == 3
b.insert(4,40)
assert b.depth() == 4
b.insert(1,10)
assert b.depth() == 4
b.insert(7,70)
assert b.depth() == 4

```

c)	**(4 Punkte)**   
Angenommen, Sie können die Schlüssel in einer selbstgewählten Reihenfolge einfügen. Welche Reihenfolge wählen Sie, damit der Baum nach dem Einfügen eine möglichst geringe Tiefe hat?

<!-- grade_id: task1c-answer, points: 4, grade: true, solution: true -->
Es wäre besser die Schlüssel in einer ungeordnete Reihenfolge einzufügen. Im Gegenfall, wenn z.B. die Schlüssel aufsteigend sortiert wären, dann fügt man immer die Schlüssel rechts in den Baum ein, damit bekommt man den Baum mit der Tiefe n, wo n die Anzahl von Elementen ist.
Z.B., nehmen wir alle Zahlen von 1 bis 10. Als erstes Element nehmen wir ein Element mit Mittelwert, z.B., 5 oder 6. Um nächstes Element links zu nehmen, könnte man wieder ein Mittelelement zwischen 0 und 6 nehmen, z.B., 3. Und so wiederholt man weiter. Die kleinste und die größte Elemente nimmt man am besten am Ende. Ein passendes Baum dafür mit der Tiefe 4, was der Formel $\left \lceil \log_{2}n \right \rceil$ entspricht, könnte so aussehen:

                6
               / \
              3   8
             / \  /\
            2  5 7  9
           /  /      \
          1  4       10

d)	**(3 Punkte)**  
Beweisen oder widerlegen Sie folgende Aussage: Wenn aus einem Suchbaum erst der Schlüssel X und danach der Schlüssel Y entfernt wird, entsteht der gleiche Baum wie bei umgekehrter Reihenfolge.

<!-- grade_id: task1d-answer, points: 3, grade: true, solution: true -->
Hier ein Gegenbeispiel:

Wir benutzen folgende Strategie:
 - Löschen: 
     - Fall.1: Der Knoten zu entfernen ist ein Blatt $\rightarrow$ einfach löschen.
     - Fall.2: Der Knoten zu entfernen hat ein Kind $\rightarrow$ Kind hochziehen.
     - Fall.3: Der Knoten zu entfernen hat zwei Kinder:
         1. Linker Teilbaum: Suche nach dem Element mit dem größten Schlüssel.
         2. Ersetze damit den Knoten zu entfernen.
         3. Entferne das unterstehende Duplikat.
         
Gegeben der Baum:
                 
                5   
               / \
              1   7
               \
                3
Wir wollen die Schlüssel 5 und 7 entfernen. Wir prüfen, ob die Reihenfolge der Löschoperationen ein unterschiedlisches Endergebnis liefert.
1. Fall: 
    - entferne 5:
                3   
               / \
              1   7
              
    - entferne 7:
                3   
               / 
              1 
2. Fall:
      - entferne 7:  
                5   
               / 
              1   
               \
                3
       - entferne 5:
                1   
                \
                 3
Also, am Ende nach zwei Löschoperationen bekommen wir zwei verschiedene Bäume, was heißt, dass die Reihenfolge eine Rolle für das Endergebnis spielt.

# Aufgabe 2 – Taschenrechner (24 Punkte)

Binärbäume eignen sich auch, um mathematische Ausdrücke auszuwerten, die als Zeichenketten der Form "2+5\*3" oder "2\*4\*(3+(4-7)\*8)-(1-6)" gegeben sind. Man bezeichnet solche Bäume als *Syntax-bäume* (https://de.wikipedia.org/wiki/Abstrakter_Syntaxbaum).  

Die Ausdrücke können die arithmetischen Operationen +, -, \*, / mit den üblichen Rechenregeln enthalten (Klammern haben die höchste Priorität, Punktrechnung geht vor Strichrechnung). Zahlen sollen der Einfachheit halber immer einstellig und positiv sein. Variablen und Funktionen kommen nicht vor. Operationen mit gleicher Priorität werden von links nach rechts ausgewertet (sogenannte *Links-Assoziativität*), damit wie gewohnt 5-2+3 = (5-2)+3 = 6 gilt (und nicht 5-(2+3) = 0, was bei Auswertung von rechts herauskäme). Steht jedoch rechts von einer Zahl oder einem Klammerausdruck eine Operation mit höherer Priorität als links, wird der rechte Operator zuerst ausgewertet. Trifft man auf eine öffnende Klammer, muss man den Substring bis zur zugehörigen schließenden Klammer suchen und die Auswertung rekursiv auf diesen Substring anwenden. Dadurch ergibt sich ein Binärbaum.  

a)	**(14 Punkte)**   
Entwickeln Sie einen Algorithmus, der den zu einem Ausdruck korrespondierenden Binärbaum aufbaut, wobei jeder innere Knoten einen Operator (+, -, \*, /) repräsentiert, jeder Unterbaum einen linken bzw. rechten Operanden, und jedes Blatt eine Zahl. Der Baum soll dann durch eine Funktion parse(s) erstellt werden, an die der Ausdruck als String übergeben wird und die den Wurzelknoten des Baums zurückgibt (die Verwendung der Python-Funktionen eval() bzw. exec ist dabei *nicht* erlaubt). Implementieren Sie diese Funktion und erklären Sie in Kommentaren, wie der Algorithmus vorgeht.   
**Hinweis:** Implementieren Sie zunächst zwei Klassen Number und Operator, die als Blattknoten bzw. innere Knoten des Syntaxbaums dienen. Number-Objekte speichern also eine Zahl, und Operator-Objekte ein Operatorsymbol sowie den linken und rechten Operanden (Unterbaum). 


<!-- grade_id: task2a-code, points: 0, grade: false, solution: true -->
```python
class NodeNum:
    def __init__(self, key, value):
        self.key = key
        self.value = value
        
        
class NodeOp:
    def __init__(self, key, value):
        self.key = key
        self.value = value
        self.left = self.right = None

        
class Number:
    def __init__(self, number):
        self.number = number
        if isinstance(number, str):
            self.number = int(number)

    def __str__(self):
        return str(self.number)

    
class Operator:
    def __init__(self, operator):
        self.operator = operator

    def __str__(self):
        return self.operator
    
    def evaluate(self, number1, number2):
        if self.operator == '-':
            return Number(number1.number - number2.number)
        if self.operator == '+':
            return Number(number1.number + number2.number)
        if self.operator == '*':
            a = Number(number1.number * number2.number)
            return a
        if self.operator == '/':
            return Number(number1.number / number2.number)
            
   
def parse(s):
    """
    This function builds the binary tree corresponding to an expression, where each internal node 
    represents an operator (+, -, *, /), each sub-tree represents a left or right operand, 
    and each leaf represents a number. 

    Arguments:
    s -- string, the input expression

    Returns:
    root -- the root node of the tree
    
    Example:
    s = '5+6*3'
    fixed = ['5','6','3','*','+']
    start: stack = []
    after token '5': stack = [ NodeNum('5', Number('5')) ]
    after token '6': stack = [ NodeNum('6', Number('6')), NodeNum('5', Number('5')) ]
    after token '3': stack = [ NodeNum('3', Number('3')), NodeNum('6', Number('6')), NodeNum('5', Number('5')) ]
    after token '*': stack = [ NodeNum('5', Number('5')), NodeOp('*', Operator('*')) ]   # NodeOp has the NodeNums '3' and '6' as children
    after token '+': stack = [ NodeOp('+', Operator('+')) ]  # NodeOp has the NodeNum '5' and NodeOp '*' as children
    -> root = stack[0] = NodeOp('+', Operator('+'))
    """
    # Strategy seen on https://stackoverflow.com/questions/423898/postfix-notation-to-expression-tree
    
    if len(s) < 3:
        raise ValueError("Expression too short.")
    if ' ' in s:
        raise ValueError("Bad syntax; no spaces please.")
        
    fixed = fix(s)        # postfix notation

    stack = []
    for token in fixed:   # start at the beginning
        
        if token.isdigit():
            stack.insert(0, NodeNum(token, Number(token)))
        
        else:          # token is an operator
            op = NodeOp(token, Operator(token))  # create new NodeOp
            
            # delete first two elements
            try:
                right = stack.pop(0)     
                left = stack.pop(0)
            except IndexError:          # a primitive syntax checker; if the stack is empty, something is wrong
                raise ValueError("Invalid mathematical expression.")
            
            # add them as operands of the new operator token
            op.right = right
            op.left = left
            stack.insert(0, op)   # push to stack; now the NodeOp can also be some other NodeOp's child

    root = stack[0]
    return root


ops = {'*':2, '/':2, '+':1, '-':1}
def fix(s):
    """
    Converts a mathematical expression from infix into postfix notation.
    Algorithm as seen on https://en.wikipedia.org/wiki/Shunting-yard_algorithm#The_algorithm_in_detail.
    
    Arguments:
    s -- string, the input expression

    Returns:
    numbers_queue -- list, tokens in postfix order
    
    """
    numbers_queue = []
    operator_stack = []
    i = 0
    while i < len(s):
        if s[i].isdigit():
            numbers_queue.append(s[i])
        if s[i] in ops:
            if operator_stack:
                while operator_stack and operator_stack[0] != "(" and ops[operator_stack[0]] >= ops[s[i]] :
                    numbers_queue.append(operator_stack.pop(0))
            operator_stack.insert(0, s[i])
        if s[i] == "(":
            operator_stack.insert(0, s[i])
        if s[i] == ")":
            try:
                while operator_stack[0] != "(":
                    numbers_queue.append(operator_stack.pop(0))
                operator_stack.pop(0)
            except IndexError:
                raise ValueError("Bad syntax, possibly missing '('.")
        i += 1
    
    while operator_stack:
        numbers_queue.append(operator_stack.pop(0))

    return numbers_queue


# Traverse
def preOrderTraverse2(node):
    if type(node) == NodeOp:
        print ("key", node.key, "value", node.value)
        preOrderTraverse2(node.left)
        preOrderTraverse2(node.right)
    elif type(node) == NodeNum:
        print ("key", node.key, "value", node.value)
            
def inOrderTraverse2(node):
    if type(node) == NodeOp:
        inOrderTraverse2(node.left)
        print ("key", node.key, "value", node.value)
        inOrderTraverse2(node.right)
    elif type(node) == NodeNum:
        print ("key", node.key, "value", node.value)

def postOrderTraverse2(node):
    if type(node) == NodeOp:
        postOrderTraverse2(node.left)
        postOrderTraverse2(node.right)
        print ("key", node.key, "value", node.value)
    elif type(node) == NodeNum:
        print ("key", node.key, "value", node.value)
```

<!-- grade_id: task2a-test, points: 12, grade: true, solution: true -->
```python
p = parse("5+6*3")
print("root:", p.key, '\n')
preOrderTraverse2(p)
print("\n")
inOrderTraverse2(p)
print("\n")
postOrderTraverse2(p)
```

Output:
```
root: + 

key + value +
key 5 value 5
key * value *
key 6 value 6
key 3 value 3


key 5 value 5
key + value +
key 6 value 6
key * value *
key 3 value 3


key 5 value 5
key 6 value 6
key 3 value 3
key * value *
key + value +
```

b)	**(3 Punkte)**   
Skizzieren Sie die Bäume, die sich für die Ausdrücke "2+5\*3" und "2\*4\*(3+(4-7)\*8)-(1-6)" ergeben.

<!-- grade_id: task2b-answer, points: 2, grade: true, solution: true -->
 2+5*3
              
               +    
             /   \        
            2     *
                 / \
                5   3 
 
 2*4*(3+(4-7)*8)-(1-6)
                                                        
                 -                                   
               /   \
              *     -
            /  \   /  \      
           *    +  1   6
          / \  / \
         2   4 3  *
                 / \
                -   8
               / \ 
              4   7

<!-- grade_id: task2b-test, points: 1, grade: true, solution: true -->
```python
import unittest
p1 = parse("2+5*3")
p2 = parse("2*4*(3+(4-7)*8)-(1-6)")

assert p1.right.key == '*'
assert p2.left.right.left.key == '3'
assert p2.left.right.right.left.key == '-'
assert p2.right.left.key == '1'

unittest_test_case = unittest.TestCase('__init__')
with unittest_test_case.assertRaises(AttributeError):    # leaf nodes have no children
    x = p2.right.right.right
with unittest_test_case.assertRaises(ValueError):        # syntax check
    parse('3*')
with unittest_test_case.assertRaises(ValueError):        # syntax check
    parse('3+5)*9')
with unittest_test_case.assertRaises(ValueError):        # syntax check
    parse('')
with unittest_test_case.assertRaises(ValueError):        # syntax check
    parse('      ')
```

c)	**(3 Punkte)**   
Implementieren Sie eine Funktion evaluateTree(root), um einen Ausdruck mit Hilfe des in a) erstellten Baums auszurechnen.

<!-- grade_id: task2c-answer, points: 0, grade: false, solution: true -->
```python
def evaluateTree(root):
    """
    This function evaluates the result of an expression given its tree root. 

    Arguments:
    root -- the root node of the tree

    Returns:
    root.value -- the result of the expression
    """
    
    if isinstance(root, NodeNum):
        return root.value
    if isinstance(root, NodeOp):
        root.value = root.value.evaluate(evaluateTree(root.left), evaluateTree(root.right))
        return root.value
    
assert evaluateTree(parse('2+5*3')).number == 17
assert evaluateTree(parse('2*4*(3+(4-7)*8)-(1-6)')).number == -163
```

<!-- grade_id: task2c-test, points: 3, grade: true, solution: false -->
```python
assert evaluateTree(parse('2+5*3')) == 17
assert evaluateTree(parse('2*4*(3+(4-7)*8)-(1-6)')) == -163

```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-5-6e161ba4ca0f> in <module>()
----> 1 assert evaluateTree(parse('2+5*3')) == 17
      2 assert evaluateTree(parse('2*4*(3+(4-7)*8)-(1-6)')) == -163

AssertionError:
```

d)	**(4 Punkte)**   
Schreiben Sie Unit Tests für Ihr Verfahren. Beachten Sie dabei die Hinweise zum Erstellen guter Tests im Kapitel "Korrektheit" des Skripts. Beispielsweise müssen Sie verschiedene Varianten der Operator-Präzedenz und Klammerung testen.

<!-- grade_id: task2d-test, points: 4, grade: true, solution: true -->
```python
import unittest
assert evaluateTree(parse('((2+(5*3)))')).number == 17      # too many + unnessecary brackets
assert evaluateTree(parse('2+3-5+3')).number == 3           # left-associativeness 
assert evaluateTree(parse('2+3*5+3')).number == 20          # preference
assert evaluateTree(parse('2+3/(3+3)')).number == 2.5       # preference
assert evaluateTree(parse('(2+3)*(5+3)')).number == 40      # preference
assert evaluateTree(parse('2*(5+2)*3')).number == 42        # preference
assert evaluateTree(parse('2*5-4+6/3')).number == 8         # preference

unittest_test_case = unittest.TestCase('__init__')
unittest_test_case.assertAlmostEqual(evaluateTree(parse('7/3')).number, 2.33, 2)    
```

```python

```