Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Introduction

In this exercise you will be learning the basics of programing with python and Jupyter (formerly iPython) Notebooks (Tasks 1 and 2) and algorithmic thinking (Tasks 3 and 4).

The exercises of this lecture will be using Jupyter in combination with the nbgrader extension. Each exercise will contain tasks that are to be answered in free text or with source code. The source code should adhere to PEP-8 style guidelines and should be commented so that the tutor understands what you are doing (points may be deducted in both cases). In some cases, a source answer cell might be followed by a Python unittest. If this test fails, you might not get points. On the other hand, if it succeeds you might not either (in case you try to cheat the test):

```python
	def pythagoras(a, b):
		return 5

	assert pythagoras(3, 4) == 5
```

# Introduction to Jupyter Notebooks and Python (7 Pts.)
 
> The Jupyter Notebook is an open-source web application that allows you to create and share documents that contain live code, equations, visualizations and narrative text. Uses include: data cleaning and transformation, numerical simulation, statistical modeling, data visualization, machine learning, and much more. - https://jupyter.org/index.html

1. Install Python 3.5 or higher with Jupyter (<http://jupyter.org/install>)
	- Suggestion for Windows: Use Anaconda
	- Suggestion for Linux-OSes: Use a virtual environment and pip
1. (Optional but highly recommended) Work through chapters 3, 4, 5, 7.2, 8 und 9.3 to get a first look at python. Chapters 4.7, 5.1.3, 5.1.4, 5.6, 5.7 and 5.8 may be skipped. Furthermore, have a look at <docs.python.org/library/index.html> to get an overview of built-in packages. Use this notebook as a scratch pad by inserting a new cell below.



## Answer the following questions
(Answer in the markdown cell, refer to code in the following code cell.)
1. In which module can the function `sqrt()` be found?
1. What happens when `sqrt()` is called with a negative number?
1. Implement two functions `mysqrt()` that prints *`mysqrt()` does not work with negative numbers, you fool.* when it is called with a negative number, and the result in the other case.
    * Use `if: … else:`
    * Use `try: … except:`
1. Implement a loop that prints `i mod 5 = <result>` for `i = [-10..10]` and explain the result of the modulo operation.
1. When should a string be enclosed in triple quotes?
1. What is the difference between the classes `dict` and `list`?
1. What is the purpose of the `__init__()` function of a class and how is it used?


<!-- grade_id: cell-1f1f5d99b4ea02f9, points: 7, grade: true, solution: true -->
1) In the `math` module <br>
2) `ValueError` is raised <br>
3) Code below

```python
from math import sqrt
```

<!-- grade_id: cell-aa57d5b55494c658, points: 7, grade: true, solution: true -->
```python
def mysqrt(n):
    if n < 0:
        print("mysqrt() does not work with negative numbers, you fool")
    else:
        print(m.sqrt(n))

```

```python
def mysqrt(n):
    try:
        print(m.sqrt(n))
    except ValueError: 
        print("mysqrt() does not work with negative numbers, you fool")

```

4) The modulo operation returns the division rest, e.g. 7 : 5 = 1 (rest 2), so 7 mod 5 = 2  

Code below

```python
for i in range(-10, 11):
    print(i, 'mod 5 =', i % 5)

```

Output:
```
-10 mod 5 = 0
-9 mod 5 = 1
-8 mod 5 = 2
-7 mod 5 = 3
-6 mod 5 = 4
-5 mod 5 = 0
-4 mod 5 = 1
-3 mod 5 = 2
-2 mod 5 = 3
-1 mod 5 = 4
0 mod 5 = 0
1 mod 5 = 1
2 mod 5 = 2
3 mod 5 = 3
4 mod 5 = 4
5 mod 5 = 0
6 mod 5 = 1
7 mod 5 = 2
8 mod 5 = 3
9 mod 5 = 4
10 mod 5 = 0
```

5) A triple quoted string after a function or class definition is its docstring. help(function_name) returns this docstring (if it exists). Triple quoted strings can also contain line breaks.

6) Both are data types in Python. 

`dict`  `d = {'a': 1, 2: 'b', (3, 4): 'b'}`
   - consists of key:value pairs, where the key allows access to the value: `d['a'] -> 1` 
   - each key can exist only once
   - unordered
   - keys have to be from an immutable data type (integers, floats, tuples, strings, but not lists)
   
`list`    `li = [12, 'a', 3.5, ['lala', 3]]`
   - is a list of elements, each element is accessed by its index: `li[0] -> 12`
   - an element can occur multiple times in a list
   - ordered
   - mutable and immutable data types are allowed
   
7) Function definition:  def __init__(self, par1, par2):

  `__init__` is the class constructor in Python. It is called implicitly through a class variable assignment, e.g. m = MyClass(). `__init__`, like all non-static class methods, always receives the current object as an implicit first parameter (here `self`. Further parameters are optional, and can be used to initialize class attributes.
  
  Usage example below.

```python
class MyClass:
    
    def __init__(self, name, number):
        self.name = name
        self.number = number
    
    def __str__(self):
        return '{} has {} apples.'.format(self.name, self.number)
        
m = MyClass('Annie', 4)
print(m)
```

Output:
```
Annie has 4 apples.
```

# Sieve of Eratosthenes (6 Pts.)
Implement the [Sieve](https://en.wiktionary.org/wiki/sieve) of Eratosthenes used to find prime numbers (To be found on the internet). Make the tests pass.

```python
sieve(n)  # print all primes below n.
```

<!-- grade_id: sieve, points: 0, grade: false, solution: true -->
```python
from math import sqrt, ceil


def sieve(n):
    """Compute and return all prime numbers below n+1"""

    if n <= 0:
        raise ValueError("There are no negative prime numbers.")
    else:
        primes = list(range(2, n + 1))   # we assume all numbers are primes

        for i in range(2, ceil(sqrt(n)) + 1):  # and remove the non-primes
            primes = [elm for elm in primes if elm % i != 0 or elm == i]

        return primes
    
```

<!-- grade_id: test_sieve_values, points: 3, grade: true, solution: false -->
```python
assert sieve(2) == [2]
assert sieve(3) == [2, 3]
assert sieve(4) == [2, 3]
assert sieve(5) == [2, 3, 5]

```

<!-- grade_id: test_sieve_invalid_input, points: 3, grade: true, solution: false -->
```python
try:
    sieve(0)
except ValueError:
    pass
else:
    raise AssertionError("Did not raise")
    
try:
    sieve(-4)
except ValueError:
    pass
else:
    raise AssertionError("Did not raise")

```

# Friday 13th (8 Pts.)

Once every few month there will be a Friday 13th, but how often does this event actually occur?

1. How many possible weekday - date combinations are possible? (Don’t forget leap years)
1. Determine how many Friday 13th you have experienced in your life


<!-- grade_id: Antwort1, points: 8, grade: true, solution: true -->
YOUR ANSWER HERE

# Stable Glasses (12 Pts.)
Glases in pubs are prone to breaking. Therefore, more stable glasses are being developed. To evaluate the stability of a new kind of glas, the following experiment is being done: Platforms that allow a controlled fall of a glass can be attached to a 3 m high stand with possible platform positions every 10 cm. The goal of this experiment is to find the highest possible fall height for a glass not to shatter with as few tries as possible. One try is defined as one fall of a glass off a platform. Another constraint is the number of available glasses which means only a certain amount of glasses may break.

1. How do you determine the height of fall in case only one glass is at hand? How many tries are needed at most (explain). From an algorithmic standpoint, this is the worst case, for pubs the best case.
1. It is possible to use fewer tries to reach the goal if more glasses may break. What is the process in this case, how many tries are needed at most and how many glasses are broken at most?
1. A compromise is to allow two glasses to break. Which process minimizes the number of tries in the worst case scenario in this test case and how many tries will be needed in the worst case? (Hint: less than 15)
1. Generalize your result from 3. s.t. the stand has `n` possible platform position. The process remains the same but the number of tries in the worst case is now a function `worst_case(n)`. Determine the function! Since the the number of tries is aways an integer, the function is a step function: The results for certain intervals may be the same until it increases by one only to be constant again for a few more steps. For `n=30`, the result should be the same as in 3.

<!-- grade_id: Antwort2, points: 10, grade: true, solution: true -->
1. Strategy
   - Start at the lowest platform (10 cm) and increase the height of the platform by 10 cm if the glass doesn't break. If it breaks, the maximum height is found.
   - The maximal number of tries is 30, in the case of the glass falling from the last platform (3 m) and remaining whole.

2. Strategy
   - Start at the middle height platform (150 cm).
   - Push a glass off the platform. If the glass breaks, lower the height to the halfway point to the bottom. If it doesn't break, increase the platform height to halfway to the top, and so on. (If there is no platform position at this middle point height, use the platform position directly underneath.)

   Essentially, this strategy produces a binary decision tree (see graph below), where each node represents a try at a specific platform height. Left edge represents a failed experiment (glass broke), right edge means glass remained whole. 

                                                  150
                                                 /    \
                                               70      220
                                              /  \     /  \
                                            30   110  ...  ...
                                           /  \    \
                                          10   50   ...
                                              /  \
                                             ...  ...
                                             
   Because it is a binary tree, we can calculate its depth, which should give us the number of tries to find the maximum          height at which the glass remains whole. 

   Depth of a binary tree = $  \log_{2}\#leaves  $    
   => We would need $\log_{2}30 = 5$ tries.  (31 leaves if we include the 0cm height as possible solution)
   
   After drawing the entire tree, it becomes clear that (only) in the case of 300cm being the maximal height, 6 tries are needed to determine that. The worst-case number of broken glasses is 4.
   
3. Strategy
    - Set 5 markers, at equal distances along the stand (50cm, 100cm, 150cm, and so on).
    - Start at the first marker (50cm).
    - Push down the first glass. If it remains whole, go up to the next marker. If it breaks, move the platform 4 positions down (e.g. 50cm -> 10cm), and use the second glass to find the maximal height, using the strategy in 1. Repeat this step until the problem is solved.
     
     In the worst case (where the wanted height is 290cm), the approach would need 10 tries to find that maximal height.

4. Assuming the platform positions are still at fixed intervals...

<!-- grade_id: cell-a2426d9a7636bfde, points: 0, grade: false, solution: true -->
```python
from math import ceil, sqrt, floor

def worst_case(n):
    """Compute worst case."""
    
    raise NotImplementedError()

```

<!-- grade_id: cell-f80d4f8640c3d48e, points: 2, grade: true, solution: false -->
```python
try:
    worst_case(0)
except ValueError:
    pass
else:
    raise AssertionError("Did not raise")

```

Output:
```
<ipython-input-11-af0f5eeb1083> in <module>()
      1 try:
----> 2     worst_case(0)
      3 except ValueError:
      4     pass
      5 else:

<ipython-input-10-d00033c35580> in worst_case(n)
      4     """Compute worst case."""
      5 
----> 6     raise NotImplementedError()

NotImplementedError:
```

```python

```