Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

#  Einleitung
In dieser Aufgabe soll Ihr Verständnis von Sortieralgorithmen gestärkt werden. Stellen Sie sicher, dass sie das Archiv samt Test-Dateien heruntergeladen und in einen Unterordner `test_cases` extrahiert haben (Quelle: http://blog.quibb.org/2008/11/sort-optimization/). Öffnen Sie den Ordner und schauen Sie sich den Schema an, nachdem die Dateien benannt sind.




# Sortierverfahren implementieren (6⋅4 Pkt.)
Implementieren Sie folgende Sortieralgorithmen.

- Selectionsort
- Insertionsort
- Quicksort
- Bubblesort
- Mergesort
- Heapsort

Die Sortieralgorithmen sollen eine sortierte Kopie der ursprünglichen Liste zurück geben und das Original nicht verändern (damit das Verhalten kongruent zu `sorted()`) ist.


<!-- grade_id: cell-6d812715dbe509ce, points: 0, grade: false, solution: true -->
```python
def selectionsort(data):
    #data = data[:]
    for i in range(len(data) - 1):
        least = i
        for j in range(i + 1, len(data)):
            if data[least] > data[j]:
                least = j
        data[least], data[i] = data[i], data[least]
    return data


def insertionsort(data):
    data = data[:]
    for i in range(1, len(data)):
        current_value = data[i]
        position = i
        while position > 0 and data[position - 1] > current_value:
            data[position] = data[position - 1]
            position -= 1
        data[position] = current_value
    return data


def bubblesort(data):
    data = data[:]
    n = len(data)
    for j in range(n - 1):
        for i in range(n - j - 1):
            if data[i] > data[i + 1]:
                data[i], data[i + 1] = data[i + 1], data[i]
    return data


def quicksort(data):
    # YOUR CODE HERE
    raise NotImplementedError()
    

def make_min_heap(data):
    heap = []

    for elm in data:
        heap.append(elm)
        i = len(heap) - 1     
        while i > 0:            # maintain heap property
            parent = (i - 1) // 2
            if heap[i] < heap[parent]:
                heap[i], heap[parent] = heap[parent], heap[i]
            i = parent

    return heap


def heapsort(data):

    heap = make_min_heap(data)
    sorted_data = []
    
    while heap:
        heap[0], heap[-1] = heap[-1], heap[0]
        sorted_data.append(heap.pop())
        heap = make_min_heap(heap)

    return sorted_data


def merge(left, right):

    merged = []
    i = 0
    j = 0
    
    while i < len(left) and j < len(right):
        if left[i] < right[j]:
            merged.append(left[i])
            i += 1
        else:
            merged.append(right[j])
            j += 1

    while i < len(left):
        merged.append(left[i])
        i += 1

    while j < len(right):
        merged.append(right[j])
        j += 1

    return merged


def mergesort(data):

    if len(data) > 1:
        middle = len(data) // 2
        left_half = data[:middle]
        right_half = data[middle:]

        sorted_left = mergesort(left_half)
        sorted_right = mergesort(right_half)

        return merge(sorted_left, sorted_right)

    else:
        return data

```

<!-- grade_id: cell-848f5caf61f83589, points: 0, grade: true, solution: false -->
```python
import random

SAMPLE_LIST = random.sample(range(100), 30)

def test_algorithm(algorithm):
    sample_list = SAMPLE_LIST[:]
    print(sample_list)
    print(algorithm(sample_list))
    assert algorithm(sample_list) == sorted(sample_list), "{} did not sort the list".format(algorithm.__name__)
    assert sample_list == SAMPLE_LIST, "{} altered the original".format(algorithm.__name__)
```

<!-- grade_id: cell-87b11519e17f1c63, points: 4, grade: true, solution: false -->
```python
test_algorithm(selectionsort)
```

Output:
```
[9, 90, 99, 18, 93, 0, 26, 15, 61, 77, 94, 4, 7, 44, 57, 88, 3, 81, 21, 23, 58, 34, 53, 72, 20, 40, 60, 65, 43, 78]
[0, 3, 4, 7, 9, 15, 18, 20, 21, 23, 26, 34, 40, 43, 44, 53, 57, 58, 60, 61, 65, 72, 77, 78, 81, 88, 90, 93, 94, 99]
<ipython-input-3-aac1f048cc98> in <module>()
----> 1 test_algorithm(selectionsort)

<ipython-input-2-b001aabae34a> in test_algorithm(algorithm)
      8     print(algorithm(sample_list))
      9     assert algorithm(sample_list) == sorted(sample_list), "{} did not sort the list".format(algorithm.__name__)
---> 10     assert sample_list == SAMPLE_LIST, "{} altered the original".format(algorithm.__name__)

AssertionError: selectionsort altered the original
```

<!-- grade_id: cell-24cfee0a571db291, points: 4, grade: true, solution: false -->
```python
test_algorithm(insertionsort)
```

Output:
```
[5, 89, 60, 39, 71, 96, 6, 10, 17, 29, 73, 95, 99, 27, 57, 11, 18, 46, 67, 92, 32, 70, 80, 98, 41, 91, 64, 20, 51, 26]
[5, 6, 10, 11, 17, 18, 20, 26, 27, 29, 32, 39, 41, 46, 51, 57, 60, 64, 67, 70, 71, 73, 80, 89, 91, 92, 95, 96, 98, 99]
```

<!-- grade_id: cell-78ffe6f52e47fc47, points: 4, grade: true, solution: false -->
```python
test_algorithm(bubblesort)
```

Output:
```
[5, 89, 60, 39, 71, 96, 6, 10, 17, 29, 73, 95, 99, 27, 57, 11, 18, 46, 67, 92, 32, 70, 80, 98, 41, 91, 64, 20, 51, 26]
[5, 6, 10, 11, 17, 18, 20, 26, 27, 29, 32, 39, 41, 46, 51, 57, 60, 64, 67, 70, 71, 73, 80, 89, 91, 92, 95, 96, 98, 99]
```

<!-- grade_id: cell-e60d068bfdb9cdde, points: 4, grade: true, solution: false -->
```python
test_algorithm(quicksort)
```

Output:
```
[5, 89, 60, 39, 71, 96, 6, 10, 17, 29, 73, 95, 99, 27, 57, 11, 18, 46, 67, 92, 32, 70, 80, 98, 41, 91, 64, 20, 51, 26]
<ipython-input-35-b001aabae34a> in test_algorithm(algorithm)
      6     sample_list = SAMPLE_LIST[:]
      7     print(sample_list)
----> 8     print(algorithm(sample_list))
      9     assert algorithm(sample_list) == sorted(sample_list), "{} did not sort the list".format(algorithm.__name__)
     10     assert sample_list == SAMPLE_LIST, "{} altered the original".format(algorithm.__name__)

<ipython-input-34-016c2ce329f1> in quicksort(data)
     34 def quicksort(data):
     35     # YOUR CODE HERE
---> 36     raise NotImplementedError()
     37 
     38 

NotImplementedError:
```

<!-- grade_id: cell-037297fb44664ec1, points: 4, grade: true, solution: false -->
```python
test_algorithm(heapsort)
```

Output:
```
[5, 89, 60, 39, 71, 96, 6, 10, 17, 29, 73, 95, 99, 27, 57, 11, 18, 46, 67, 92, 32, 70, 80, 98, 41, 91, 64, 20, 51, 26]
[5, 6, 10, 11, 17, 18, 20, 26, 27, 29, 32, 39, 41, 46, 51, 57, 60, 64, 67, 70, 71, 73, 80, 89, 91, 92, 95, 96, 98, 99]
```

<!-- grade_id: cell-96e6149464b940df, points: 4, grade: true, solution: false -->
```python
test_algorithm(mergesort)
```

Output:
```
[42, 21, 37, 72, 30, 85, 28, 79, 56, 60, 48, 58, 50, 39, 4, 17, 40, 70, 24, 52, 49, 89, 18, 65, 57, 86, 95, 19, 11, 83]
[4, 11, 17, 18, 19, 21, 24, 28, 30, 37, 39, 40, 42, 48, 49, 50, 52, 56, 57, 58, 60, 65, 70, 72, 79, 83, 85, 86, 89, 95]
```

# Einlesen von Dateien (4 Pkt.)
Schreiben sie eine Funktion `load_test_cases(path)`, die zu den zur Verfügung gestellten Testfällen jeweils die Datei mit 10e3 Einträgen lädt.

<!-- grade_id: cell-d64d799097780639, points: 0, grade: false, solution: true -->
```python
import os

TEST_CASE_PATH = "test_files"

def load_test_cases(path):
    """
    Args:
    path string
    
    Returns:
    test_cases dictionary 
    """
    test_cases = {}
    path = os.path.join(path, '')    # adds trailing slash

    for entry in os.scandir(path):
        if entry.name.endswith(".10000.txt"):
            test_cases[entry.name.replace(".10000.txt", "")] = [int(x) for x in open(path + entry.name).read().strip().split('\n')]
                                                          
    return test_cases

```

<!-- grade_id: cell-6b64b68014341d9a, points: 4, grade: true, solution: false -->
```python
import unittest

unittest_test_case = unittest.TestCase('__init__')

loaded_test_cases = load_test_cases(TEST_CASE_PATH)
available_test_cases = ['1245.repeat', 'end.0', 'identical', 'med.3.killer', 'rand.dups', 'rand.no.dups', 'rand.steps', 
                        'rev.ends', 'rev.partial', 'rev.saw', 'reverse', 'seq.0.is.1000', 'seq.partial', 'seq.saw', 'sequential']

unittest_test_case.assertCountEqual(loaded_test_cases.keys(), available_test_cases)

for key, value in loaded_test_cases.items():
    assert len(value) == 10e3, "Number of testcases mismatch for {}".format(key)
```

# Sortierverfahren testen (12 Pkt.)
 Vergleichen Sie die Sortierverfahren untereinander und mit der Python Funktion `sorted()`
 - Benutzen sie `result = %timeit -n 100 -r 3 -o` um die Ausführung zu messen und das Ergebnis in einer Variablen speichern zu können. Auf das beste Ergebnis können Sie mit `result.best` zugreifen.
 - Gruppieren Sie die Ergebnisse nach Testfall.
 - Testen sie folgende „Splices“ Ihres Testdatensatzes: `x_values = [10, 20, 50, 100, 200, 500, 1000, 2000, 5000, 10000]`
 - Stellen sie die Ergebnisse mittels Tabelle und Graphen mit Titel und Achsenbeschriftung dar (Matplotlib).
 - Bewerten Sie das Ergebnis

<!-- grade_id: cell-584a475079bfa4bd, points: 12, grade: true, solution: true -->
```python
%matplotlib
import matplotlib.pyplot as plt
# for quicksort
import sys
sys.setrecursionlimit(20000)


def get_times(algorithm, testset):
    times = []
    for val in x_values:
        result = %timeit -n 10 -r 3 -o algorithm(testset[:val])
        times.append(result.best)
    return times

x_values = [10, 20, 50, 100, 200, 500, 1000, 2000, 5000]


'''
plt.plot(x_values, get_times(selectionsort))
plt.plot(x_values, get_times(insertionsort))
plt.plot(x_values, get_times(bubblesort))
#plt.plot(x_values, get_times(quicksort))
plt.plot(x_values, get_times(heapsort))
plt.plot(x_values, get_times(mergesort))
plt.plot(x_values, get_times(sorted))
'''
#plt.plot(x_values, get_times(selectionsort, loaded_test_cases['rand.dups']))
#plt.plot(x_values, get_times(insertionsort, loaded_test_cases['rand.dups']))
#plt.plot(x_values, get_times(bubblesort, loaded_test_cases['rand.dups']))
#plt.plot(x_values, get_times(quicksort))
plt.plot(x_values, get_times(heapsort, loaded_test_cases['rand.dups']))
#plt.plot(x_values, get_times(mergesort, loaded_test_cases['rand.dups']))
#plt.plot(x_values, get_times(sorted, loaded_test_cases['rand.dups']))

plt.title("Running times for 'rand.dups")
plt.xlabel("Array Size")
plt.ylabel("Seconds")
plt.xticks(x_values)
plt.legend(['selectionsort', 'insertionsort', 'bubblesort', 'heapsort', 'mergesort', 'sorted'], loc='upper left')

plt.show()

```

Output:
```
Using matplotlib backend: TkAgg
69.9 µs ± 2.59 µs per loop (mean ± std. dev. of 3 runs, 10 loops each)
328 µs ± 25.1 µs per loop (mean ± std. dev. of 3 runs, 10 loops each)
2.04 ms ± 166 µs per loop (mean ± std. dev. of 3 runs, 10 loops each)
9.39 ms ± 811 µs per loop (mean ± std. dev. of 3 runs, 10 loops each)
41.1 ms ± 380 µs per loop (mean ± std. dev. of 3 runs, 10 loops each)
316 ms ± 19.1 ms per loop (mean ± std. dev. of 3 runs, 10 loops each)
1.39 s ± 20.3 ms per loop (mean ± std. dev. of 3 runs, 10 loops each)
6.1 s ± 5.66 ms per loop (mean ± std. dev. of 3 runs, 10 loops each)
<ipython-input-34-016c2ce329f1> in heapsort(data)
     60         heap[0], heap[-1] = heap[-1], heap[0]
     61         sorted_data.append(heap.pop())
---> 62         heap = make_min_heap(heap)
     63 
     64     return sorted_data

<ipython-input-34-016c2ce329f1> in make_min_heap(data)
     44         i = len(heap) - 1
     45         while i > 0:            # maintain heap property
---> 46            
```

Bewertung:

Naja, alles super.

```python

```