Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Aufgabe 1 - Eigenschaften der Sortierung (8 Punkte)

a)	Um Daten sortieren zu können, muss auf den Elementen eine Relation '≤' definiert sein, die die Eigenschaften einer totalen Ordnung (siehe https://de.wikipedia.org/wiki/Ordnungsrelation) hat. Zeigen Sie, dass dadurch auch die anderen Relationen '<', '>', '≥', '==' und '!=' festgelegt sind, indem Sie diese so definieren, dass in den Definitionen nur die '≤' Relation und logische Operationen (not, and, or) vorkommen.

<!-- grade_id: answer1a, points: 0, grade: false, solution: true -->
```python
def is_less_than_or_equal_to(a, b): 
    if a <= b:
        return True
    else:
        return False


def is_less_than(a, b):
    """Define '<' with '<=', not, and, or"""
    if not b <= a: 
        return True
    else:
        return False


def is_more_than(a, b):
    """Define '>' with '<=', not, and, or"""
    if not a <= b:
        return True
    else:
        return False


def is_more_than_or_equal_to(a, b):
    """Define '>=' with '<=', not, and, or"""
    if b <= a :
        return True
    else:
        return False


def is_equal_to(a,b):
    """Define '==' with '<=', not, and, or"""
    if a <= b and b <= a:
        return True
    else:
        return False
    

def is_not_equal_to(a,b):
    """Define '!=' with '<=', not, and, or"""
    if not a <= b or not b <= a:
        return True
    else:
        return False

```

<!-- grade_id: test1a, points: 6, grade: true, solution: false -->
```python
assert is_less_than(1,2) == True
assert is_more_than(2,1) == True
assert is_more_than_or_equal_to(2,1) == True
assert is_equal_to(1,1) == True
assert is_not_equal_to(1,2) == True

```

b)	Nach dem Sortieren eines Arrays a der Größe n ist das folgende Axiom erfüllt: Für alle Paare von Indizes j und k mit j < k sind die zugehörigen Arrayelemente sortiert, d.h. es gilt a[j] ≤ a[k]. Das sind n(n-1)/2 Bedingungen (eine für jedes mögliche Paar j < k), die man alle prüfen muss, um die Korrektheit der Sortierung nachzuweisen. Zeigen Sie mit Hilfe der Eigenschaften der '≤' Relation, dass dieser Aufwand unnötig ist und man auch mit (n-1) Tests auskommt. Welche Tests sollte man ausführen, und warum sind die übrigen dann automatisch erfüllt?

<!-- grade_id: answer1b, points: 2, grade: true, solution: true -->
Um zu prüfen, ob ein Array sortiert ist, reicht es, nur benachbarte Elemente schrittweise miteinander zu vergleichen. Die Relation '≤' ist nämlich transitiv und reflexiv, also a[i] ≤ a[k] gdw. a[i] ≤ a[j] and a[j] ≤ a[k] (i < j < k).

Z.B., für ein Array [1, 2, 3, 4, 5] (n = 5) ist der Aufwand n - 1 = 4 (geprüft wird 1 ≤ 2, 2 ≤ 3 usw.)

# Aufgabe 2 – Zelluläre Automaten (18 Punkte)

Algorithmen werden wesentlich durch die Menge der elementaren Operationen geprägt, aus denen sie aufgebaut sind. Zelluläre Automaten sind interessant, weil sie aus sehr einfachen elementaren Operationen erstaunlich komplexes Verhalten generieren. Das bekannteste Beispiel ist Conway’s Game of Life (http://de.wikipedia.org/wiki/Conways_Spiel_des_Lebens).

Jeder zelluläre Automat besteht aus unabhängigen Zellen, die auf einem regelmäßigen Gitter an-geordnet sind und K vorher festgelegte Zustände annehmen können. Die Zustände werden meist durch druckbare Zeichen wie "␣" (Leerzeichen), "\*", "#" etc. symbolisiert. Der Automat entwickelt sich in diskreten Zeitschritten, und der Zustand jeder Zelle zum Zeitpunkt T+1 ergibt sich deterministisch aus den Zuständen dieser Zelle sowie ihrer unmittelbaren Nachbarn zum Zeitpunkt T. Die Menge der erlaubten Zustandsübergänge in Abhängigkeit der möglichen Nachbarschafts-Konstellationen wird als Regel des Automaten bezeichnet. Die Regel wird einmal festgelegt und bleibt dann für alle Zeitschritte gleich.

1. **(4 Punkte)** Wir beschränken uns in dieser Übung auf 1-dimensionale zelluläre Automaten. Diese kann man bequem als Strings einer vorher festgelegten Länge darstellen, sodass jedes Zeichen im String eine Zelle repräsentiert. Der neue Zustand einer Zelle j ergibt sich dann aus dem alten Zustand dieser Zelle und dem Zustand ihrer linken und rechten Nachbarzelle (Zellen j-1 und j+1), also jeweils aus Teilstrings der Länge 3. Deshalb kann man die Regel bequem in einem Python-Dictionary darstellen: Strings der Länge 3 bilden die Schlüssel, und die zugehörigen Werte geben den neuen Zustand der mittleren Zelle an, z.B.:  
rule = { 
         "␣␣␣" : "␣",    # drei Leerzeichen => Leerzeichen
         "␣*␣" : "*",   # Stern zwischen zwei Leerzeichen => Stern
         ...           # weitere Zustandsübergänge
         }
Implementieren Sie eine Funktion  
*ca2 = ca_step(ca1, rule)*  
der der String ca1 (Zustand zur Zeit T) sowie die Regel übergeben werden und die daraus den neuen String ca2 (mit gleicher Länge wie ca1) als Zustand zur Zeit T+1 berechnet. Definieren Sie dabei eine geeignete Randbehandlung für die beiden Enden von ca1, also für die Zellen, die keinen linken bzw. rechten Nachbarn haben. Diese Funktion wird in b) und c) benutzt.


<!-- grade_id: answer2a, points: 0, grade: false, solution: true -->
```python
def ca_step(ca1, rule):
    """
    The string ca1 (state at time T) and the rule are passed
    and from this the new string ca2 (with the same length as ca1) is calculated as state at time T + 1. 
    Define a suitable edge treatment for the two ends of ca1, i.e. for the cells that have no left or right neighbors.

    Arguments:
    ca1 -- string
    rule -- Python dictionary

    Returns:
    ca2 -- string, with the same length as ca1
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: test2a, points: 4, grade: true, solution: false -->
```python
# DEFINE THE RULE (DO NOT CHANGE THIS!)
rule = {"   " : "#",
        "*  " : " ",
        " * " : " ",
        "  *" : " ",
        "#  " : "*",
        " # " : "*",
        "  #" : "*",
        "***" : "#",
        " **" : "*",
        "* *" : "*",
        "** " : "*",
        "#**" : " ",
        "*#*" : " ",
        "**#" : " ",
        "###" : "#",
        " ##" : " ",
        "# #" : " ",
        "## " : " ",
        "*##" : "*",
        "#*#" : "*",
        "##*" : "*",
        " *#" : "#",
        " #*" : "#",
        "* #" : "#",
        "*# " : "#",
        "# *" : "#",
        "#* " : "#"}
        
assert ca_step('#*#*#', rule) == '** **'
assert ca_step('# **#  ***', rule) == ' #* #* *##'

```

2.**(6 Punkte)**  
Elementare zelluläre Automaten haben nur zwei Zustände "␣" (Leerzeichen), Zustand "\*". Wie viele verschiedene Regeln kann man daraus aufbauen? Manche dieser Regeln führen zu interessantem Verhalten. Experimentieren Sie mit den Regeldefinitionen (indem Sie die Dictionaries variieren) und berechnen Sie jeweils die ersten 30 Zeitschritte, wobei der Anfangszustand immer ein String der Länge 71 ist, bei dem sich nur das Element 35 im Zustand "\*" befindet, alle übrigen sind im Zustand "␣". Geben Sie mit print den aktuellen Zustand nach jeder Iteration aus. Implementieren Sie vier Funktionen "elementary1()" bis "elementary4()" (die jeweils eine Regeldefinition und die Schleife enthalten) die nette Muster erzeugen.

<!-- grade_id: answer2b1, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

<!-- grade_id: answer2b2, points: 0, grade: false, solution: true -->
```python
def elementary1(initial_states, time_steps):
    """
    Contain a rule definition and the loop for printing the states at each time step.

    Arguments:
    initial_states -- string
    time_steps -- scalar
    """
    # YOUR RULE
    rule = {"   " : " ",
            "*  " : "*",
            " * " : " ",
            "  *" : "*",
            "***" : " ",
            " **" : "*",
            "* *" : " ",
            "** " : "*"}
    
    # YOUR CODE HERE
    raise NotImplementedError()
         
def elementary2(initial_states, time_steps):
    """
    Contain a rule definition and the loop for printing the states at each time step.

    Arguments:
    initial_states -- string
    time_steps -- scalar
    """
    # YOUR CODE HERE
    raise NotImplementedError()
        
def elementary3(initial_states, time_steps):
    """
    Contain a rule definition and the loop for printing the states at each time step.

    Arguments:
    initial_states -- string
    time_steps -- scalar
    """
    # YOUR CODE HERE
    raise NotImplementedError()
    
def elementary4(initial_states, time_steps):
    """
    Contain a rule definition and the loop for printing the states at each time step.

    Arguments:
    initial_states -- string
    time_steps -- scalar
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

Test the four functions one-by-one:

<!-- grade_id: test2b1, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: test2b2, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: test2b3, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: test2b4, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

3.**(8 Punkte)**  
Implementieren Sie jetzt eine Funktion pingpong(), deren Regel eine Art Ping-Pong-Spiel realisiert. Es gibt hier neben dem Zustand "␣" einen Zustand "#", der die beiden Spielfeldränder angibt, sowie einen Ball, der immer zwischen den beiden Rändern hin- und herfliegt und au-tomatisch reflektiert wird. Der Ball wird durch zwei benachbarte Zellen symbolisiert, die die Zustände "o-" haben, wenn der Ball gerade nach links fliegt, und "-o" wenn er nach rechts fliegt. Der Anfangszustand ist der String "␣#␣␣␣␣␣␣␣␣o-␣␣#␣" (der Abstand zwischen den beiden "#" beträgt 12 Zeichen und soll unverändert bleiben). Implementieren Sie die zeitliche Entwicklung des Automaten als Endlosschleife (diese kann man in Python durch die Tasten-kombination "Ctrl-C" abbrechen, wenn man lange genug zugeschaut hat). Um die Animation besser beobachten zu können, sollten Sie in jeder Iteration eine gewisse Wartezeit einbauen (Funktion "sleep()" aus dem Modul "time"). Beachten Sie, dass die Regel während des Ablaufs fix bleibt und nicht verändert werden darf (um z.B. die Flugrichtung des Balls umzukehren).


<!-- grade_id: answer2c, points: 0, grade: false, solution: true -->
```python
import time

def pingpong(initial_states, time_steps):
    """
    Contain a ping-pong game rule definition and the loop for printing the states at each time step.

    Arguments:
    initial_states -- string
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

Test your pingpong function with an initial states " #    o-      # " for 30 time steps:

<!-- grade_id: test2c, points: 8, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

# Aufgabe 3 – Datenstruktur für Array, Stack und Queue (14 Punkte)

Wir haben in der Vorlesung die abstrakten Containertypen *Array*, *Stack* und *Queue* kennengelernt. In dieser Aufgabe sollen Sie eine Python-Klasse UniversalContainer implementieren und testen, die die Fähigkeiten aller drei Container vereinigt. Die folgende Funktionalität soll unterstützt werden:
        
        c = UniversalContainer()   # create empty container  
        c.size()                   # get number of elements  
        c.capacity()               # get number of available memory cells  
        c.push(v)                  # append element v at the end  
        c.popFirst()               # remove first element (Queue functionality)  
        c.popLast()                # remove last element (Stack functionality)              
        v = c[k]                   # read element at index k (Array functionality)  
        c[k] = v                   # replace element at k (Array functionality)  
        v = c.first()              # read first element (Queue functionality)  
        v = c.last()               # read last element (Stack functionality)  

Die Klasse hat zu diesem Zweck einen internen Speicherbereich data\_ der Größe capacity\_, der aber nur bis zur Größe size\_ gefüllt ist. Beim Aufruf c.push(v) wird Element v in die nächste freie Speicherzelle von data\_ geschrieben und size\_ inkrementiert. Gilt allerdings vor dem push, dass size\_ == capacity\_, muss man erst einen Speicherbereich mit größerer (z.B. verdoppelter) capacity\_  schaffen und die vorhandenen Daten dahin umkopieren. Bei popLast() wird einfach size\_ dekrementiert, so dass das vorletzte Element zum neuen letzten Element wird. Bei popFirst() werden alle Elemente einen Index heruntergeschoben und dadurch das erste überschrieben.

Der folgende Code gibt den Rahmen der Klasse vor, den Sie nur noch vervollständigen müssen:

<!-- grade_id: answer3, points: 0, grade: false, solution: true -->
```python
class UniversalContainer:
    
    # constructor for empty container
    def __init__(self):                      
        self.capacity_ = 1                   # we reserve memory for at least one item
        self.data_ = [None]*self.capacity_   # the internal memory
        self.size_ = 0                       # no item has been inserted yet
    
    def size(self):
        return self.size_
    
    def capacity(self):
        # YOUR CODE HERE
        raise NotImplementedError()
    
    # add item at the end
    def push(self, item):
        # YOUR CODE HERE
        raise NotImplementedError()
    
    # remove first element
    def popFirst(self):
        # YOUR CODE HERE
        raise NotImplementedError()
        
    # remove last element
    def popLast(self):
        # YOUR CODE HERE
        raise NotImplementedError()
        
    # __getitem__ implements v = c[index]
    def __getitem__(self, index):     
        if index < 0 or index >= self.size_:
            raise RuntimeError("index out of range")
        # YOUR CODE HERE
        raise NotImplementedError()
    
    # __setitem__ implements c[index] = v
    def __setitem__(self, index, v):  
        if index < 0 or index >= self.size_:
            raise RuntimeError("index out of range")
        # YOUR CODE HERE
        raise NotImplementedError()
    
    # read first element
    def first(self):
        # YOUR CODE HERE
        raise NotImplementedError()
    
    # read last element
    def last(self):
        # YOUR CODE HERE
        raise NotImplementedError()
```

Wir haben in der Vorlesung gelernt, dass die Datenstruktur folgende Axiome erfüllt:
+ Ein neuer Container hat die Größe 0.  
+ Zu jeder Zeit gilt: c.size() <= c.capacity().  
+ Nach einem push gilt: (i) die Größe hat sich um eins erhöht, (ii) das gerade eingefügte Element ist jetzt das letzte, (iii) alle anderen Elemente haben sich nicht verändert, (iv) wenn der Container vorher leer war, ist das eingefügte Element auch das erste, andernfalls bleibt das erste Element unverändert.  
+ Nach c[k] = v gilt: (i) die Größe bleibt unverändert, (ii) Index k enthält das Element v, (iii) die übrigen Elemente haben sich nicht verändert.  
+ Nach c.popLast() gilt: (i) die Größe hat sich um eins verringert, (ii) die Elemente vom ersten bis zum vorletzten bleiben unverändert.  
+ Nach c.popFirst() gilt: (i) die Größe hat sich um eins verringert, (ii) das zweite bis letzte Element sind einen Index nach unten gerückt.  
+ Wenn der Container nicht leer ist, gilt stets (i) c.first() == c[0] und (ii) c.last() == c[c.size()-1].  

Aufgaben:  
a)	Vervollständigen Sie die Implementierung von UniversalContainer.  
b)	Schreiben Sie eine Funktion testContainer(), die für verschiedene repräsentative Fälle mittels assert (z.B. 'assert c.size() == 0') testet, dass die obigen Axiome gelten.


<!-- grade_id: test3a, points: 14, grade: true, solution: false -->
```python
c = UniversalContainer()
assert c.size() == 0
assert c.capacity() == 1
c.push(1)
c.push(2)
c.push(3)
c.push(4)
c.push(5)
assert c.size() == 5
assert c.capacity() == 8
assert c.data_[0] == 1
assert c.data_[1] == 2
assert c.data_[2] == 3
assert c.data_[3] == 4
assert c.data_[4] == 5
assert c.data_[5] == None
assert c.data_[6] == None
assert c.data_[7] == None
c.popFirst()
assert c.data_[0] == 2
assert c.data_[1] == 3
assert c.data_[2] == 4
assert c.data_[3] == 5
assert c.data_[4] == None
assert c.data_[5] == None
assert c.data_[6] == None
assert c.data_[7] == None
assert c.size() == 4
c.popLast()
assert c.data_[0] == 2
assert c.data_[1] == 3
assert c.data_[2] == 4
assert c.data_[3] == None
assert c.data_[4] == None
assert c.data_[5] == None
assert c.data_[6] == None
assert c.data_[7] == None
assert c.size() == 3
assert c.__getitem__(0) == 2
assert c.__getitem__(1) == 3
assert c.__getitem__(2) == 4
c.__setitem__(0,100)
assert c.first() == 100
assert c.last() == 4
```