Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Maryna Charniuk"
COLLABORATORS = "Dung Nguyen, Lyubomira Dimitrova"
```

---

# HASHING

Schreiben Sie eine einfache Hash-Funktion `my_hash(string, s)` mit der Tabellengröße $s$ mittels Modulo Operation (auch bekannt als *Lineares Sondieren* oder  *division remainder hashing*. (4 Punkte)

<!-- grade_id: cell-af1b70e92b19d8c4, points: 4, grade: true, solution: true -->
```python
def my_hash(string, s):
    sum_chars = 0
    for char in string:
        sum_chars += ord(char)
    return sum_chars % s
    
```

Berechnen Sie mittels der von Ihnen geschriebene Funhtion die Hashwerte der Zeichenketten **done** und **node** mit einer Tabellengröße von $s = 3$. (3 Punkte)

<!-- grade_id: cell-b27ef90cc2132399, points: 3, grade: true, solution: true -->
```python
print(my_hash('done', 3))
print(my_hash('node', 3))

assert my_hash('done', 3) == my_hash('node', 3)
```

Output:
```
2
2
```

Ist der Hash beider Zeichenketten der Gleiche? Nennen Sie den Grund. (3. Punkte)

<!-- grade_id: cell-71dfab42b9d4fe7a, points: 3, grade: true, solution: true -->
Ja, weil wir die ASCII Werten aller Zeichen im String addieren, also spielt die Reihenfolge der Zeichen keine Rolle. Um das zu vermeiden, könnte man jedem Zeichen ein Gewicht geben, z.B. die Position im String.

Eine weit verbreitete Methode zur Kollisionsvermeidung ist die *Lineare Suche (linear probing)*. Erzeugen Sie eine Hash-Tabelle der Größe $s = 11$ in der Form eines Assoziativen Datenfelds (Dictionary) für die Liste `lisf = [10,45,43,76,57,12,77,13]` Die Hash-Tabelle sollte wie folgt aussehen: `table = {0: 76, 1: 12, 2: 45, 3: 10, 4: 43, 5: 77, 6: 13, 7: None, 8: None, 9: None, 10: 57}`. (10 Punkte)

<!-- grade_id: cell-f34574fb999a906f, points: 0, grade: false, solution: true -->
```python
# first we need to modify the hash function to handle integers
def my_hash(item, s):
    if isinstance(item, int):
        return item % s
    if isinstance(item, str):
        sum_chars = 0
        for char in item:
            sum_chars += ord(char)
        return sum_chars % s


def linearProbingHash(item_list, hash_table_size):
    
    table = {i: None for i in range(hash_table_size)}
    
    for value in item_list:
        key = my_hash(value, hash_table_size)
        
        if None in table.values():     # check if there are any free spaces in the hash table
            while table[key]:
                if key == hash_table_size - 1:      # make table circular (if you reach the end, go to 0)
                    key = -1
                key += 1
                
        else:
            raise RuntimeError("Hash-Table too small.") 
        
        table[key] = value
    
    return table

```

<!-- grade_id: cell-9d77dc861e6c0eb0, points: 10, grade: true, solution: false -->
```python
lisf = [10,45,43,76,57,12,77,13]
print(linearProbingHash(lisf, 11))

```

Output:
```
{0: 43, 1: 45, 2: 76, 3: 57, 4: 12, 5: 77, 6: 13, 7: None, 8: None, 9: None, 10: 10}
```

Erzeugen Sie eine Hash-Tabelle der Größe $s = 11$ in Form eines Assoziativen Datenfelds für die Liste `list = ["his", "her", "this", "that", "what", "when", "how", "why", "i dont know"].`Die Hash-Tabelle sollte aussehen wie folgt: `table = {0: 'her', 1: 'this', ... , 10: None}`. (10 Punkte)

<!-- grade_id: cell-edc469c308109c3d, points: 0, grade: false, solution: true -->
```python
def linearProbingHashStrings(item_list, hash_table_size):
    
    table = {i: None for i in range(hash_table_size)}
    
    for value in item_list:
        key = my_hash(value, hash_table_size)
        
        if None in table.values():     # check if there are any free spaces in the hash table
            while table[key]:
                if key == hash_table_size - 1:      # make table circular (if you reach the end, go to 0)
                    key = -1
                key += 1
                
        else:
            raise RuntimeError("Hash-Table too small.") 
        
        table[key] = value
    
    return table

```

<!-- grade_id: cell-df52a49cc705fd10, points: 10, grade: true, solution: false -->
```python
stringList = ["his", "her", "this", "that", "what", "when", "how", "why", "i dont know"]
print(linearProbingHashStrings(stringList, 11))
#print(linearProbingHashStrings(stringList, 7))    # raises a RuntimeError because 7 < len(list)
```

Output:
```
{0: 'her', 1: 'this', 2: None, 3: 'why', 4: 'that', 5: 'his', 6: 'when', 7: 'what', 8: 'how', 9: 'i dont know', 10: None}
```

Schreiben Sie eine Funktien `all_cocktails(filename)`, die die Datei `cocktails.json` in ein Assozietives Datenfeld `recipes` liest und schreiben Sie eine Funktion `all_ingredients(recipes)`, die eine komplette Liste aller Zutaten ausgibt. (10 Punkte)

<!-- grade_id: cell-a39e1b1cae7075df, points: 0, grade: false, solution: true -->
```python
import json

def all_cocktails(filename):
    with open(filename) as f:
        j = json.load(f)
    return {i: cocktail for i, cocktail in enumerate(j['cocktails'])}
        
def all_ingredients(recipes):
    ingredients = set()
    for r in recipes.values():
        for listing in r['ingredients']:
            try:
                ingredients.add(listing['ingredient'])
            except KeyError:      # not all children of 'ingredients' contain 'ingredient'
                pass
    return ingredients


recipes = all_cocktails('cocktails.json')
numberOfIngredients = len(all_ingredients(recipes))
print (all_ingredients(recipes))
print (numberOfIngredients)

```

Output:
```
{'Cherry liqueur', 'Syrup', 'Grapefruit juice', 'Tequila', 'Cola', 'Galliano', 'Red Port', 'Dark rum', 'Peach puree', 'Lemon juice', 'Pineapple juice', 'Creme liqueur', 'Champagne', 'Lillet Blonde', 'Ginger Ale', 'Coffee liqueur', 'Prosecco', 'Campari', 'Orange juice', 'Pisco', 'Egg yolk', 'Lime juice', 'Vodka', 'Kirsch', 'Cognac', 'Triple Sec', 'Soda water', 'Apricot brandy', 'Whiskey', 'Cream', 'White rum', 'Cranberry juice', 'Coconut milk', 'Gin', 'Calvados', 'Drambuie', 'Vermouth'}
37
```

<!-- grade_id: cell-4d70aaa8919259a4, points: 10, grade: true, solution: false -->
```python
recipes = all_cocktails('cocktails.json')

assert (numberOfIngredients) == (37)
assert ('Apricot brandy' in all_ingredients(recipes)) 
assert ('Pineapple juice' in all_ingredients(recipes))
assert ('Campari' in all_ingredients(recipes)) 
assert ('Kirsch' in all_ingredients(recipes))
assert ('Pisco' in all_ingredients(recipes))
```